---
##  Problem 1: Iterator (6 pts)
Implement a custom range iterator that mimics some of the basic functionality of Python's built-in `range()` function (You can not use `range()` direclty). This exercise will help you understand how iterators work in Python.

**Note:** there is the same question as our in-class practice.

##  Problem 1.1: A custom class called `MyRange` (3 pts)
1. Create a class named `MyRange` that takes two integer arguments in its constructor:

     **start:**  the first value of the range

   
     **start:**: the value to stop before (exclusive)

3. The class should implement the iterator protocol, making it usable in a for loop.

4. Each iteration should return the next integer in the sequence, starting from start and incrementing by 1 each time.
The iteration should stop when the value would become equal to or greater than end.

<!-- grade_id: class_my_range, points: 0, grade: false, solution: true -->
```python
class MyRange:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_MYRANGE, points: 3, grade: true, solution: false -->
```python
nums = MyRange(1, 5)
assert next(nums) == 1, "First value should be 1"
assert next(nums) == 2, "Second value should be 2"
assert next(nums) == 3, "Third value should be 3"
assert next(nums) == 4, "Fourth value should be 4"

try:
    next(nums)
    assert False, "StopIteration should have been raised"
except StopIteration:
    pass

```

##  Problem 1.2: A custom generator called `my_range` (3 pts)
Create a function `my_range(start, end)` that generates a sequence of numbers from start to end (exclusive). The function should be implemented as a generator.

<!-- grade_id: class_my_range2, points: 0, grade: false, solution: true -->
```python
def my_range(start, end):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_range, points: 3, grade: true, solution: false -->
```python
nums = my_range(1, 5)
assert next(nums) == 1, "First value should be 1"
assert next(nums) == 2, "Second value should be 2"
assert next(nums) == 3, "Third value should be 3"
assert next(nums) == 4, "Fourth value should be 4"

try:
    next(nums)
    assert False, "StopIteration should have been raised"
except StopIteration:
    pass

```

##  Problem 2: Memoization (6 pts)

###  2.1: Climbing stairs (3 pts)
You are climbing a staircase. It takes n steps to reach the top.

Each time you can either climb 1 or 2 steps. In how many distinct ways can you climb to the top? Let's try not use recursion in your implementation. Instead, use an iterative approach or use memoization.

<!-- grade_id: climbStairs, points: 0, grade: false, solution: true -->
```python
def climb_stairs(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_stairs, points: 3, grade: true, solution: false -->
```python
assert climb_stairs(1) == 1, "Test case 1 failed"
assert climb_stairs(2) == 2, "Test case 2 failed"
assert climb_stairs(3) == 3, "Test case 3 failed"
assert climb_stairs(7) == 21, "Test case 4 failed"
assert climb_stairs(35) == 14930352, "Test case 5 failed (larger input)"
```

###  2.2: Two Sum(3 pts)
Given an array of integers nums and an integer `target`, return indices of the two numbers such that they add up to `target`.

You may assume that each input would have exactly one solution, and you may not use the same element twice. You can return the answer in any order.

**Note:**: This is a classical interview question. Try to use a dictionary as a "mini database" to store and quickly look up values.

<!-- grade_id: sum, points: 0, grade: false, solution: true -->
```python
def two_sum(nums, target):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_two_sum, points: 3, grade: true, solution: false -->
```python
assert two_sum([2, 7, 11, 15], 9) == [0, 1], "Test case 1 failed"
assert two_sum([3, 2, 4], 6) == [1, 2], "Test case 2 failed"
```

##  Problem 3: Error Handling (6 pts)

###  3.1: Safely Divide(3 pts)
Write a function called `safe_divide(a, b)` that safely performs division of a by b. The function should:
1) Return the result of a divided by b if b is not zero.

2) Return the string "Cannot divide by zero" if b is zero.

3) Use a try-except block to handle the `ZeroDivisionError`.


<!-- grade_id: divide, points: 0, grade: false, solution: true -->
```python
def safe_divide(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_divide, points: 3, grade: true, solution: false -->
```python
assert safe_divide(10, 2) == 5.0, "Test case 1 failed"
assert safe_divide(10, 0) == "Cannot divide by zero", "Test case 2 failed"
assert isinstance(safe_divide(10, 3), float), "Test case 5 failed: Result should be a float"
```

###  3.2: Raise our own error (3 pts)
Create a function called `safe_divide(a, b)` that does the following:

Takes an age parameter (integer).
1) If age is not int or float, raise `TypeError` with message "Age must be a number".
2) If the age is less than 0, raise a `ValueError` with message "Age cannot be negative".
3) If the age is over 150, raise the `ValueError` with the message "Age cannot be over 150".



<!-- grade_id: validate, points: 0, grade: false, solution: true -->
```python
def validate_age(age):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_error, points: 3, grade: true, solution: false -->
```python
assert validate_age(25) == "Valid age: 25", "Test for valid age failed"
try:
    validate_age("twenty")
    assert False, "Non-numeric age should raise TypeError"
except TypeError as e:
    assert str(e) == "Age must be a number", "Incorrect error message for non-numeric input"
    
try:
    validate_age(-5)
    assert False, "Negative age should raise ValueError"
except ValueError as e:
    assert str(e) == "Age cannot be negative", "Incorrect error message for negative age"

```