### Week 5-2 NumPy Practice Problems

```python
import numpy as np
```

### Warm up
NumPy creation and indexing.

```python
# 1. Creates 3x3 array of zeros
np.zeros((3,3))
```

Output:
```
array([[0., 0., 0.],
       [0., 0., 0.],
       [0., 0., 0.]])
```

```python
# 2. Creates array [0, 2, 4, 6, 8]
np.array([2*i for i in range(5)])
```

Output:
```
array([0, 2, 4, 6, 8])
```

```python
# 3. 5 evenly spaced numbers from 0 to 1
np.arange(0, 1, 0.2)
```

Output:
```
array([0. , 0.2, 0.4, 0.6, 0.8])
```

```python
# Array indexing, Gets elements > 5
arr = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
```

### Problem 1
Convert an array of Celsius measurements to Fahrenheit.
Note: `fahrenheit = (Celsius * 9/5) + 32`.

```python
def convert_to_fahrenheit_list(temperatures):
    temperatures = np.array(temperatures)
    convert = lambda temps : temps * 9 / 5 + 32
    return convert(temperatures)
```

```python
print("Testing with a regular list:")
temps_list = [0, 10, 20, 30]
print(convert_to_fahrenheit_list(temps_list))
```

Output:
```
Testing with a regular list:
[32. 50. 68. 86.]
```

```python
def convert_to_fahrenheit_numpy_array(temperatures):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
print("\nTesting with a NumPy array:")
temps_array = np.array([0, 10, 20, 30])
print(convert_to_fahrenheit_numpy_array(temps_array))
```

### Problem 2
Use numpy to create a function that adds a border filled with zeros around an existing array.

*Hint*: Use numpy's `pad` function. See https://numpy.org/doc/stable/reference/generated/numpy.pad.html

Note: we did not cover this function in class, see if you are comfortable of reading numpy documents and learn on your own!

<!-- grade_id: cell-2466342b05a17106, points: 0, grade: false, solution: true -->
```python
import numpy as np
def add_border(x):
    ans = np.pad(x, ((1,1), (1,1)), 'constant', constant_values=0)
    return ans
```

<!-- grade_id: cell-fb9bae34e3908aed, points: 1, grade: true, solution: false -->
```python
x = np.ones((1, 1))
claimed = add_border(x)
correct = np.zeros((3, 3))
correct[1, 1] = 1.0
assert claimed.shape == correct.shape
assert np.all(np.abs(claimed - correct) < 1e-8)
```

### Problem 3
Create an $n \times n$ numpy array filled with a checkerboard pattern of zeros and ones. The first entry should be the integer `0`. Assume $n \ge 2$.

<!-- grade_id: cell-9982d08b4ea7057d, points: 0, grade: false, solution: true -->
```python
def make_checkerboard(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e6f9eb1adf863ce2, points: 1, grade: true, solution: false -->
```python
correct = \
    [[0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0]]
assert np.all(correct == make_checkerboard(8))
```

```python
import numpy as np
x = np.arange(10)
```